```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet4.ipynb")
```

# Worksheet 4: Aggregations and ACID transaction in MongoDB

## Establish Pymongo connection

Recall that in worksheet1, we have created a cluster on MongoDB Atlas and connect to it via Pymongo and a credentials JSON file. Let's reestablish that connection for our exercises.

- Make sure you use the `adsc_3610` conda environment.
- You might need to copy & paste the `credentials_mongodb.json` file that you used in worksheet1 to the working directory of worksheet3.

```python
from pymongo import MongoClient # import mongo client to connect
import json # import json to load credentials
import urllib.parse

# load credentials from json file
with open('credentials_mongodb.json') as f:
    login = json.load(f)

#Added below line for format issues
import urllib.parse 

# Manually reasigning KEy values as I was getting errors
login = {
    'username': 'hassanalisyd',
    'password': 'hassanalisyd',
    'host': 'cluster0.ktepq.mongodb.net'
}

# Ensuring the key exists in the dictionary as I was getting KeyErrirs 
username = login.get('username')
password = login.get('password')
host = login.get('host')

# Print the values to debug
print(f"Username: {username}")
print(f"Password: {password}")
print(f"Host: {host}")

#Adding a check using If-else to make sure the keys work ok

if username and password and host:
    # Encode the password
    encoded_password = urllib.parse.quote(password)
    url = "mongodb+srv://{}:{}@{}/?retryWrites=true&w=majority".format(username, encoded_password, host)
    print(url)
else:
    print("One or more keys are missing in the login dictionary.")

```

Output:
```
Username: hassanalisyd
Password: hassanalisyd
Host: cluster0.ktepq.mongodb.net
mongodb+srv://hassanalisyd:hassanalisyd@cluster0.ktepq.mongodb.net/?retryWrites=true&w=majority
```

```python
# connect to the database
client = MongoClient(url)
```

```python
print(client.server_info())  # Test the connection

```

Output:
```
{'version': '7.0.12', 'gitVersion': 'b6513ce0781db6818e24619e8a461eae90bc94fc', 'modules': ['enterprise'], 'allocator': 'tcmalloc', 'javascriptEngine': 'mozjs', 'sysInfo': 'deprecated', 'versionArray': [7, 0, 12, 0], 'bits': 64, 'debug': False, 'maxBsonObjectSize': 16777216, 'storageEngines': ['devnull', 'inMemory', 'queryable_wt', 'wiredTiger'], 'ok': 1.0, '$clusterTime': {'clusterTime': Timestamp(1727311533, 15), 'signature': {'hash': b'\xbe\x07j\xf7\xaf\xfe\xc9@HV\xe5\xce\xc0\x84|i\x80\x08RI', 'keyId': 7366339589516034049}}, 'operationTime': Timestamp(1727311533, 15)}
```

```python
# drop database books and students if they exist
client.drop_database('bookstore')
client.drop_database('students')
```

To test if your connection has been succesful, let's try to print out all the databases

```python
# list all databases
client.list_database_names()
```

Output:
```
['sample_airbnb',
 'sample_analytics',
 'sample_geospatial',
 'sample_guides',
 'sample_mflix',
 'sample_restaurants',
 'sample_supplies',
 'sample_training',
 'sample_weatherdata',
 'school',
 'shop',
 'admin',
 'local']
```

- Load `sample_mflix` database into an object called db
- Load the `movies` collection into an object called `collection`

```python
db = client['sample_mflix']
collection = db['movies']
```

```python
...

```

Output:
```
Ellipsis
```

```python
grader.check("ex0")
```

Output:
```
ex0 results: All test cases passed!
```

#### Exercise 1: Calculate the Average Rating of Movies by Genre

rubric: {accuracy = 2}

**Task**: Write an aggregation pipeline to calculate the average imdb rating of movies grouped by genre.

**Hint**: First, use `$unwind` stage to unwind the `genres` array before grouping by genre. Use the `$group` stage to group by genre and the `$avg` operator to calculate the average rating.



Expected output should look something like this:

```json
{'_id': 'War', 'averageRating': 7.128591954022989}
{'_id': 'News', 'averageRating': 7.252272727272728}
{'_id': 'Romance', 'averageRating': 6.6564272782136396}
...
{'_id': 'Sci-Fi', 'averageRating': 6.123609653725079}
{'_id': 'Film-Noir', 'averageRating': 7.397402597402598}
```

```python
# Exercise 1: Calculate the Average Rating of Movies by Genre

pipeline1 = [

     # Unwind the genres array to have one genre per document
    {"$unwind": "$genres"},
    
    # Group by genre and calculate the average IMDb rating
    {
        "$group": {
            "_id": "$genres",  # Group by the genre field
            "averageRating": {"$avg": "$imdb.rating"}  # Calculate the average IMDb rating
        }
    },
    

]

# Execute the pipeline
result1 = list(collection.aggregate(pipeline1))

for doc in result1:
    print(doc)
```

Output:
```
{'_id': 'Talk-Show', 'averageRating': 7.0}
{'_id': 'Comedy', 'averageRating': 6.450214658080344}
{'_id': 'Thriller', 'averageRating': 6.304498977505112}
{'_id': 'Documentary', 'averageRating': 7.365679824561403}
{'_id': 'Sport', 'averageRating': 6.749041095890411}
{'_id': 'Animation', 'averageRating': 6.89669603524229}
{'_id': 'Mystery', 'averageRating': 6.527425044091711}
{'_id': 'Horror', 'averageRating': 5.784709897610922}
{'_id': 'Western', 'averageRating': 6.823553719008264}
{'_id': 'History', 'averageRating': 7.1696100917431185}
{'_id': 'Action', 'averageRating': 6.347098402018503}
{'_id': 'War', 'averageRating': 7.128591954022989}
{'_id': 'Biography', 'averageRating': 7.087984189723319}
{'_id': 'Adventure', 'averageRating': 6.493680884676145}
{'_id': 'Drama', 'averageRating': 6.803377338624768}
{'_id': 'Fantasy', 'averageRating': 6.3829847908745245}
{'_id': 'Film-Noir', 'averageRating': 7.397402597402598}
{'_id': 'Short', 'averageRating': 7.377574370709382}
{'_id': 'Musical', 'a
```

```python
grader.check("ex1")
```

Output:
```
ex1 results: All test cases passed!
```



#### Exercise 2: Find the Top 5 Movies with the Highest IMDb Rating

rubric: {accuracy = 3}

**Task**: Write an aggregation pipeline to find the top 5 movies with the highest Tomatoes Viewer rating. In the result, only keep the movies `title` and `tomatoes.viewer.rating` field only, remove `_id` from the result.

**Hint**: Use the `$sort` stage to sort by Tomatoes Viewer rating in descending order and the `$limit` stage to limit the results to the top 5. Use `$project` to select fields that you want to include/exclude in the output.



Expected output:

```json
...
{'title': 'Murder by Natural Causes', 'tomatoes': {'viewer': {'rating': 5.0}}}
{'title': 'I Am Maria', 'tomatoes': {'viewer': {'rating': 5.0}}}
```

```python
# print out tomatoes viewer rating of some movies
for doc in collection.find({}, {"_id": 0,
                                "title": 1, 
                                "tomatoes.viewer.rating": 1}).limit(5):
    print(doc)
```

Output:
```
{'title': 'The Great Train Robbery', 'tomatoes': {'viewer': {'rating': 3.7}}}
{'title': 'A Corner in Wheat', 'tomatoes': {'viewer': {'rating': 3.6}}}
{'title': 'Winsor McCay, the Famous Cartoonist of the N.Y. Herald and His Moving Comics', 'tomatoes': {'viewer': {'rating': 3.4}}}
{'title': 'Traffic in Souls', 'tomatoes': {'viewer': {'rating': 3.0}}}
{'title': 'Gertie the Dinosaur', 'tomatoes': {'viewer': {'rating': 3.7}}}
```

```python
# Exercise 2: Find the Top 5 Movies with the Highest Tomatoes Viewer Rating

pipeline = [
    # Step 1: Project the necessary fields
    {"$project": {
        "title": 1,
        "tomatoes_viewer_rating": "$tomatoes.viewer.rating"
    }},
    
    # Step 2: Filter out documents with missing or null ratings
    {"$match": {
        "tomatoes_viewer_rating": {"$ne": None}
    }},
    
    # Step 3: Sort by Tomatoes Viewer rating in descending order
    {"$sort": {
        "tomatoes_viewer_rating": -1
    }},
    
    # Step 4: Limit the result to the top 5 movies
    {"$limit": 5},
    
    # Step 5: Exclude the _id field from the output
    {"$project": {
        "_id": 0,
        "title": 1,
        "tomatoes": {
            "viewer": {
                "rating": "$tomatoes_viewer_rating"
            }
        }
    }}
]

# Execute the pipeline
result = list(collection.aggregate(pipeline))

# Print the results
for doc in result:
    print(doc)

```

Output:
```
{'title': 'Kadin Hamlet', 'tomatoes': {'viewer': {'rating': 5.0}}}
{'title': 'The Seine Meets Paris', 'tomatoes': {'viewer': {'rating': 5.0}}}
{'title': 'Night Function', 'tomatoes': {'viewer': {'rating': 5.0}}}
{'title': 'Murder by Natural Causes', 'tomatoes': {'viewer': {'rating': 5.0}}}
{'title': 'I Am Maria', 'tomatoes': {'viewer': {'rating': 5.0}}}
```

```python
grader.check("ex2")
```

Output:
```
ex2 results:
    ex2 - 1 result:
        ❌ Test case failed
        Trying:
            assert len(result2) == 5, f'Expected 5 documents, but found {len(result2)}'
        Expecting nothing
        **********************************************************************
        Line 1, in ex2 0
        Failed example:
            assert len(result2) == 5, f'Expected 5 documents, but found {len(result2)}'
        Exception raised:
            Traceback (most recent call last):
              File "c:\Users\Hassa\miniforge3\envs\adsc_3610\Lib\doctest.py", line 1368, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest ex2 0[0]>", line 1, in <module>
                assert len(result2) == 5, f'Expected 5 documents, but found {len(result2)}'
                           ^^^^^^^
            NameError: name 'result2' is not defined. Did you mean: 'result1'?
        Trying:
            assert set(result2[0].keys()) == set(['title', 'tomatoes']), f"
```



#### Exercise 3: Count the Number of Movies by Year

rubric: {accuracy = 2}

**Task**: Write an aggregation pipeline to count the number of movies released each year. Sort by count in descending order

**Hint**: Use the `$group` stage to group by year and the `$count` operator to count the number of movies. Use the `sort` stage to sort by count



Expected output:
```json
...
{'_id': 2008, 'count': 886}
{'_id': 2010, 'count': 866}
...
```

```python
# Exercise 3: Count the Number of Movies by Year


pipeline3 = [
    # Step 1: Group by year and count the number of movies for each year
    {
        "$group": {
            "_id": "$year",  # Group by the year field
            "count": {"$sum": 1}  # Count the number of movies, use 'count' as the field name
        }
    },
    
    # Step 2: Sort by the movie count in descending order
    {
        "$sort": {
            "count": -1  # Sort by count in descending order
        }
    }
]




# Execute the pipeline
result3 = list(collection.aggregate(pipeline3))
for doc in result3:
    print(doc)
```

Output:
```
{'_id': 2013, 'count': 1105}
{'_id': 2014, 'count': 1073}
{'_id': 2012, 'count': 955}
{'_id': 2009, 'count': 917}
{'_id': 2011, 'count': 893}
{'_id': 2008, 'count': 886}
{'_id': 2010, 'count': 866}
{'_id': 2007, 'count': 810}
{'_id': 2006, 'count': 774}
{'_id': 2005, 'count': 713}
{'_id': 2004, 'count': 678}
{'_id': 2002, 'count': 622}
{'_id': 2001, 'count': 612}
{'_id': 2003, 'count': 603}
{'_id': 2000, 'count': 581}
{'_id': 1999, 'count': 515}
{'_id': 1998, 'count': 513}
{'_id': 2015, 'count': 480}
{'_id': 1997, 'count': 439}
{'_id': 1996, 'count': 407}
{'_id': 1995, 'count': 372}
{'_id': 1994, 'count': 305}
{'_id': 1993, 'count': 274}
{'_id': 1992, 'count': 270}
{'_id': 1988, 'count': 251}
{'_id': 1991, 'count': 238}
{'_id': 1989, 'count': 232}
{'_id': 1990, 'count': 225}
{'_id': 1987, 'count': 222}
{'_id': 1984, 'count': 199}
{'_id': 1986, 'count': 190}
{'_id': 1985, 'count': 189}
{'_id': 1982, 'count': 177}
{'_id': 1981, 'count': 168}
{'_id': 1980, 'count': 167}
{'_id': 1983, 'cou
```

```python
grader.check("ex3")
```

Output:
```
ex3 results: All test cases passed!
```



#### Exercise 4: Find the Top 3 Directors with the Most Movies

rubric: {accuracy = 4}

**Task**: Write an aggregation pipeline to find the top 3 directors who have directed the most movies.

**Hint**: Use the `$group` stage to group by director, the `$count` operator to count the number of movies, the `$sort` stage to sort by the count in descending order, and the `$limit` stage to limit the results to the top 3.

You might need to `$unwind` the directors array.



Expected output:
```json
{'_id': ..., 'movieCount': ...}
{'_id': ..., 'movieCount': ...}
{'_id': ..., 'movieCount': ...}
```

```python
# Exercise 4: Find the Top 3 Directors with the Most Movies



pipeline4 = [
    # Step 1: Unwind the directors array (if a movie has multiple directors)
    {"$unwind": "$directors"},
    
    # Step 2: Group by director and count the number of movies for each
    {
        "$group": {
            "_id": "$directors",  # Group by the director's name
            "movieCount": {"$sum": 1}  # Count the number of movies
        }
    },
    
    # Step 3: Sort by the movie count in descending order
    {
        "$sort": {
            "movieCount": -1  # Sort by movie count in descending order
        }
    },
    
    # Step 4: Limit to the top 3 directors
    {"$limit": 3}


    
]

# Execute the pipeline
result4 = list(collection.aggregate(pipeline4))
for doc in result4:
    print(doc)
```

Output:
```
{'_id': 'Woody Allen', 'movieCount': 40}
{'_id': 'Martin Scorsese', 'movieCount': 32}
{'_id': 'Takashi Miike', 'movieCount': 31}
```

```python
grader.check("ex4")
```

Output:
```
ex4 results: All test cases passed!
```

### Exercise 5: Calculate the Average IMDb Rating and Number of Movies for Each Genre (Rounded to 2 Decimals)

rubric: {accuracy = 3}

**Task**: Write an aggregation pipeline to calculate the average IMDb rating (rounded to 2 decimals) and the number of movies for each genre.

**Hint**: Use the `$unwind` stage to deconstruct the genres array, the `$group` stage to group by genre, the `$avg` operator to calculate the average rating, and the `$project` stage to round the average rating to 2 decimal places.



Expected output:

```json
...
{'_id': 'Crime', 'movieCount': 2457, 'averageRating': 6.69}
{'_id': 'Sci-Fi', 'movieCount': 958, 'averageRating': 6.12}
...
```

```python
# Exercise 5: Calculate the Average IMDb Rating and Number of Movies for Each Genre (Rounded to 2 Decimals)



pipeline5 = [
    # Step 1: Unwind the genres array
    {"$unwind": "$genres"},
    
    # Step 2: Group by genre and calculate the movie count and average IMDb rating
    {
        "$group": {
            "_id": "$genres",  # Group by genre
            "movieCount": {"$sum": 1},  # Count the number of movies for each genre
            "averageRating": {"$avg": "$imdb.rating"}  # Calculate the average IMDb rating for each genre
        }
    },
    
    # Step 3: Project the results, rounding the average rating to 2 decimal places
    {
        "$project": {
            "_id": 1,  # Keep the genre
            "movieCount": 1,  # Keep the movie count
            "averageRating": {"$round": ["$averageRating", 2]}  # Round the average rating to 2 decimal places
        }
    },
    
    # Step 4: Optionally, sort by movie count or average rating
    {
        "$sort": {
            "movieCount": -1  # Sort by movie count in descending order (can also sort by averageRating if needed)
        }
    }

    
]

# Execute the pipeline
result5 = list(collection.aggregate(pipeline5))
for doc in result5:
    print(doc)
```

Output:
```
{'_id': 'Drama', 'movieCount': 12385, 'averageRating': 6.8}
{'_id': 'Comedy', 'movieCount': 6532, 'averageRating': 6.45}
{'_id': 'Romance', 'movieCount': 3318, 'averageRating': 6.66}
{'_id': 'Crime', 'movieCount': 2457, 'averageRating': 6.69}
{'_id': 'Thriller', 'movieCount': 2454, 'averageRating': 6.3}
{'_id': 'Action', 'movieCount': 2381, 'averageRating': 6.35}
{'_id': 'Adventure', 'movieCount': 1900, 'averageRating': 6.49}
{'_id': 'Documentary', 'movieCount': 1834, 'averageRating': 7.37}
{'_id': 'Horror', 'movieCount': 1470, 'averageRating': 5.78}
{'_id': 'Biography', 'movieCount': 1269, 'averageRating': 7.09}
{'_id': 'Family', 'movieCount': 1249, 'averageRating': 6.33}
{'_id': 'Mystery', 'movieCount': 1139, 'averageRating': 6.53}
{'_id': 'Fantasy', 'movieCount': 1055, 'averageRating': 6.38}
{'_id': 'Sci-Fi', 'movieCount': 958, 'averageRating': 6.12}
{'_id': 'Animation', 'movieCount': 912, 'averageRating': 6.9}
{'_id': 'History', 'movieCount': 874, 'averageRating': 7.17}
{'_id': 'Mu
```

```python
grader.check("ex5")
```

Output:
```
ex5 results: All test cases passed!
```

### Exercise 6: Implementing ACID Transactions for a Shopping Cart

rubric: {accuracy = 4}

In this exercise, you will implement an ACID (Atomicity, Consistency, Isolation, Durability) transaction using MongoDB. The goal is to ensure that a series of operations either all succeed or all fail, maintaining the integrity of the database.

#### Instructions

1. **Setup MongoDB Client**: Ensure you have the MongoDB client installed and connected to your database.
2. **Create Collections**: Create two collections in your MongoDB database: `inventory` and `carts`.
3. **Insert Sample Data**:
    - Insert sample items into the `inventory` collection with fields `item_id` and `quantity`.
    - Insert a sample user into the `carts` collection with fields `user_id` and `items`.
4. **Implement the Transaction**:
    - Write a function `add_to_cart` that performs the following operations within a transaction:
        - Check if the inventory has enough quantity for the specified item.
        - Deduct the quantity from the inventory.
        - Add the item to the user's shopping cart.
    - Ensure that if any operation fails, the transaction is aborted and no changes are made to the database.
5. **Run the Transaction**: Execute the transaction and handle any errors that may occur.




```python
from pymongo.errors import ConnectionFailure, OperationFailure
from pymongo.write_concern import WriteConcern
from pymongo.read_concern import ReadConcern
from pymongo.read_preferences import ReadPreference

client.drop_database('shop')  # drop the shop database if it exists

# Select the database and collections
db = client['shop']
inventory = db['inventory']
carts = db['carts']

# Insert sample data into the inventory collection
inventory.insert_many([
    {'item_id': 'item456', 'quantity': 10},
    {'item_id': 'item789', 'quantity': 5},
    {'item_id': 'item123', 'quantity': 20}
])

# Insert sample data into the carts collection
carts.insert_one({
    'user_id': 'user123',
    'items': []
})

```

Output:
```
InsertOneResult(ObjectId('66f4a0d8e7bacdc69771689d'), acknowledged=True)
```

#### 6.1 Write a function to check if the inventory has enough quantity for a particular item
- raise ValueError if the order_quantity exceed the inventory quantity
- Print out a message `Item {item_id} has enough quantity` if the order_quantity <= the inventory quantity

```python
def check_quantity(item_id, order_quantity):
    
    # Find the item in the inventory by item_id
    item = inventory.find_one({"item_id": item_id})
    
    # Check if the item exists in the inventory
    if not item:
        raise ValueError(f"Item {item_id} not found in inventory")

    # Get the available quantity of the item
    available_quantity = item['quantity']
    
    # Check if the order quantity exceeds the available quantity
    if order_quantity > available_quantity:
        raise ValueError(f"Order quantity {order_quantity} exceeds available quantity {available_quantity} for item {item_id}")
    
    # If sufficient quantity is available, print a success message
    print(f"Item {item_id} has enough quantity")




```

```python
# print out the quantity of item123 from the inventory collection
print(inventory.find_one({'item_id': 'item123'}))

```

Output:
```
{'_id': ObjectId('66f4a0d8e7bacdc69771689c'), 'item_id': 'item123', 'quantity': 20}
```

```python
check_quantity('item123', 10) # this should print out a message that item123 has enough quantity
```

Output:
```
Item item123 has enough quantity
```

```python
check_quantity('item123', 100) # this should raise a ValueError
```

Output:
```
Cell In[25], line 1
----> 1 check_quantity('item123', 100) # this should raise a ValueError

Cell In[22], line 15, in check_quantity(item_id, order_quantity)
     13 # Check if the order quantity exceeds the available quantity
     14 if order_quantity > available_quantity:
---> 15     raise ValueError(f"Order quantity {order_quantity} exceeds available quantity {available_quantity} for item {item_id}")
     17 # If sufficient quantity is available, print a success message
     18 print(f"Item {item_id} has enough quantity")

ValueError: Order quantity 100 exceeds available quantity 20 for item item123
```

#### 6.2 Write a function to update the quantity from the inventory

- Use the `update_one` function to update the quantity of inventory
- You can use the `$inc` to increment the quantity of inventory based on the order quantity

```python
def update_quantity(order_quantity, item_id):
    
    # Use the $inc operator to decrement the inventory quantity
    result = inventory.update_one(
        {"item_id": item_id},  # Find the document with the matching item_id
        {"$inc": {"quantity": -order_quantity}}  # Decrease the quantity by order_quantity
    )
    
    # Check if the update was successful
    if result.modified_count == 0:
        raise ValueError(f"Item {item_id} not found or no update made")
    
    # Print a confirmation message
    print(f"Inventory updated for item {item_id}, reduced by {order_quantity}")




```

```python
# print out the quantity of item_id 'item123' before updating
print(inventory.find_one({'item_id': 'item123'}))
```

Output:
```
{'_id': ObjectId('66f4a0d8e7bacdc69771689c'), 'item_id': 'item123', 'quantity': 20}
```

```python

# Update the quantity for item_id 'item123'
order_quantity = 10
item_id = 'item123'
update_quantity(order_quantity, item_id)
```

Output:
```
Inventory updated for item item123, reduced by 10
```

```python
# print out the quantity of item_id 'item123' after updating
print(inventory.find_one({'item_id': 'item123'}))
```

Output:
```
{'_id': ObjectId('66f3427467a7018c189676bc'), 'item_id': 'item123', 'quantity': 10}
```

#### 6.3 Write a function to update user's shopping cart

- Use the `update_one` function to update the shopping cart
- Use the `$push` operator to an element to an array. If the array does not exist, it will be created.

```python
# print out shopping cart before adding items
print(carts.find_one({'user_id': 'user123'}))
```

Output:
```
{'_id': ObjectId('66f3427467a7018c189676bd'), 'user_id': 'user123', 'items': []}
```

```python
# write a function to add items to the shopping cart
def add_to_cart(user_id, item_id, quantity):
    
    
    result = carts.update_one(
        {"user_id": user_id},  # Find the cart by user_id
        {
            "$push": {
                "items": {
                    "item_id": item_id,
                    "quantity": quantity
                }
            }
        }
    )
    
    # Check if the cart was found and updated
    if result.matched_count == 0:
        raise ValueError(f"No cart found for user {user_id}")
    
    # Print confirmation message
    print(f"Added {quantity} of {item_id} to the cart for user {user_id}")

```

```python
# test the function by adding items to the shopping cart
add_to_cart('user123', 'item123', 5)
```

Output:
```
Added 5 of item123 to the cart for user user123
```

```python
# print out shopping cart after adding items
print(carts.find_one({'user_id': 'user123'}))
```

Output:
```
{'_id': ObjectId('66f3427467a7018c189676bd'), 'user_id': 'user123', 'items': [{'item_id': 'item123', 'quantity': 5}]}
```

#### 6.4 Combine everything into a transaction function

Create a transaction function that:
- First check if the quantity is enough in the inventory, if not, it should raise a ValueError
- Update the quantity from the inventory by substracting the order_quantity
- Update user's shopping cart with the item_id and the order_quantity

Note: You need to add a line `session=session` in the update_one function for inventory and shopping carts

```python

def check_quantity(item_id, order_quantity, session):
    # Check if the inventory has enough quantity
    item = inventory.find_one({"item_id": item_id}, session=session)
    if item is None:
        raise ValueError(f"Item {item_id} not found in inventory")
    if item['quantity'] < order_quantity:
        raise ValueError(f"Not enough quantity for item {item_id}. Available: {item['quantity']}, Requested: {order_quantity}")
    print(f"Item {item_id} has enough quantity.")

def update_quantity(order_quantity, item_id, session):
    # Deduct the quantity from the inventory using $inc
    result = inventory.update_one(
        {"item_id": item_id},
        {"$inc": {"quantity": -order_quantity}},
        session=session
    )
    if result.matched_count == 0:
        raise ValueError(f"Item {item_id} not found in inventory")
    print(f"Updated inventory for item {item_id}, deducted {order_quantity} units.")

def add_item_to_cart(user_id, item_id, quantity, session):
    # Add the item to the user's shopping cart using $push
    result = carts.update_one(
        {"user_id": user_id},
        {"$push": {"items": {"item_id": item_id, "quantity": quantity}}},
        session=session
    )
    if result.matched_count == 0:
        raise ValueError(f"No cart found for user {user_id}")
    print(f"Added {quantity} of {item_id} to the cart for user {user_id}")

# Define the transaction function
def add_to_cart(session, user_id, item_id, quantity):
    # Check if the inventory has enough quantity
    # Deduct the quantity from the inventory 
    # Add the item to the user's shopping cart

    try:
        # 1. Check if the inventory has enough quantity
        check_quantity(item_id, quantity, session)
        
        # 2. Deduct the quantity from the inventory 
        update_quantity(quantity, item_id, session)
        
        # 3. Add the item to the user's shopping cart
        add_item_to_cart(user_id, item_id, quantity, session)
        
    except (ConnectionFailure, OperationFailure, ValueError) as e:
        print(f"Transaction aborted: {str(e)}")
        raise  # Re-raise the error for external handling
          
    ...

```

```python

# Start a session
with client.start_session() as session:
    try:
        # Start a transaction
        session.start_transaction()

        # Perform the add to cart operation
        add_to_cart(session, 'user123', 'item456', 2)

        # Commit the transaction
        session.commit_transaction()
        print("Transaction committed.")
    except (ConnectionFailure, OperationFailure, ValueError) as e:
        # Abort the transaction on error
        session.abort_transaction()
        print(f"Transaction aborted due to error: {e}")
```

Output:
```
Item item456 has enough quantity.
Updated inventory for item item456, deducted 2 units.
Added 2 of item456 to the cart for user user123
Transaction committed.
```

Let's test the constraint of the transaction. This should throws an error because order_quantity of 200 is exceeding the inventory quantity

```python

# Start a session
with client.start_session() as session:
    try:
        # Start a transaction
        session.start_transaction()

        # Perform the add to cart operation
        add_to_cart(session, 'user123', 'item456', 200)

        # Commit the transaction
        session.commit_transaction()
        print("Transaction committed.")
    except (ConnectionFailure, OperationFailure, ValueError) as e:
        # Abort the transaction on error
        session.abort_transaction()
        print(f"Transaction aborted due to error: {e}")
```

Output:
```
Transaction aborted: Not enough quantity for item item456. Available: 8, Requested: 200
Transaction aborted due to error: Not enough quantity for item item456. Available: 8, Requested: 200
```

**There are no auto-tests for Q6, it will be graded manually**

## Submission instructions

{rubric: mechanics = 5}

- Make sure the notebook can run from top to bottom without any error. Restart the kernel and run all cells.
- Commit and push your notebook to the github repo
- Double check your notebook is rendered properly on Github and you can see all the outputs clearly

Tests for ex0 (points: 1)
```
>>> assert db.name == 'sample_mflix', f"expected db name to be 'sample_mflix' but got {db.name}"
>>> assert set(db.list_collection_names()) == set(['sessions', 'comments', 'movies', 'theaters', 'users', 'embedded_movies'])
>>> assert collection.name == 'movies', f"expected collection name to be 'movies' but got {collection.name}"
>>> assert collection.count_documents({}) == 21349, f'expected 21349 documents but got {collection.count_documents({})}'
```


Tests for ex1 (points: 5)
```
>>> import numpy as np
>>> mystery_rating = list(filter(lambda x: x['_id'] == 'Mystery', result1))[0]['averageRating']
>>> assert np.isclose(mystery_rating, 6.527425044091, atol=0.01), f'expected average rating close to 6.53 but got {mystery_rating}'
```


Tests for ex2 (points: 5)
```
>>> assert len(result2) == 5, f'Expected 5 documents, but found {len(result2)}'
>>> assert set(result2[0].keys()) == set(['title', 'tomatoes']), f"Expected keys: ['title', 'tomatoes'], but found {result2[0].keys()}"
>>> expected_titles = ['Kadin Hamlet', 'The Seine Meets Paris', 'Night Function', 'Murder by Natural Causes', 'I Am Maria']
>>> assert [doc['title'] for doc in result2] == expected_titles, f'Expected titles: {expected_titles}, but found: {[doc['title'] for doc in result2]}'
```


Tests for ex3 (points: 5)
```
>>> count_2013 = list(filter(lambda x: x['_id'] == 2013, result3))[0]['count']
>>> assert count_2013 == 1105, f'expected count 1105 for the year 2013 but got {count_2013}'
```
```
>>> count_2014 = list(filter(lambda x: x['_id'] == 2014, result3))[0]['count']
>>> assert count_2014 == 1073, f'expected count 1073 for the year 2014 but got {count_2014}'
```


Tests for ex4 (points: 5)
```
>>> assert set([doc['_id'] for doc in result4]) == set(['Woody Allen', 'Martin Scorsese', 'Takashi Miike']), f"expected directors: ['Woody Allen', 'Martin Scorsese', 'Takashi Miike'], but got {[doc['_id'] for doc in result4]}"
```
```
>>> assert list(filter(lambda x: x['_id'] == 'Woody Allen', result4))[0]['movieCount'] == 40, f'expected Woody Allen to have 40 movies, but got {list(filter(lambda x: x['_id'] == 'Woody Allen', result4))[0]['movieCount']}'
```


Tests for ex5 (points: 5)
```
>>> adventure_doc = list(filter(lambda x: x['_id'] == 'Adventure', result5))[0]
>>> assert adventure_doc['movieCount'] == 1900, f"expected movieCount 1900 for 'Adventure' but got {adventure_doc['movieCount']}"
>>> assert adventure_doc['averageRating'] == 6.49, f"expected averageRating 6.49 for 'Adventure' but got {adventure_doc['averageRating']}"
```
